```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Homework 04: Food Safety (Continued)

You must submit this assignment to Gradescope by the on-time deadline. **We strongly encourage you to plan to submit your work to Gradescope several days (hours) before the stated deadline.** This way, you will have ample time to reach out to staff for support if you encounter difficulties with submission. While course staff is happy to help guide you with submitting your assignment ahead of the deadline, course staff may not be available to assist with **last-minute submission issues**.

**Please read the instructions carefully when you are submitting your work to Gradescope**.


## This Assignment

In this homework, we will continue our exploration of restaurant food safety scores for restaurants in San Francisco. The main goal for this assignment is to focus more on the analysis of the dataset, building on the data cleaning we have done earlier in HW 03. 

## Ethical Note
- This dataset contains publicly available health inspection scores for restaurants in San Francisco.  
- While the data is used here for educational purposes, please interpret findings responsibly and avoid drawing unfair conclusions about specific businesses.




## Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.

Finally, unless we state otherwise, **do not use for loops or list comprehensions**. The majority of this assignment can be done using built-in commands in `pandas` and `NumPy`.  Our autograder isn't smart enough to check, but you're depriving yourself of key learning objectives if you write loops / comprehensions, and you also won't be ready for the midterm.

### Debugging Guide
If you run into any technical issues, we highly recommend checking out the [Debugging Guide](https://mtu.instructure.com/courses/1619791/pages/debugging-guide). In this guide, you can find general questions about Jupyter notebooks / Datahub, Gradescope, and common pandas errors.

```python
import numpy as np
import pandas as pd

import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
sns.set()
plt.style.use('fivethirtyeight')
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

In HW02, we took you through the entire process of reading data from a file to perform some exploration of the data. Here, we again load the dataset that we will be using in HW04 along with some of the columns we had added in HW02. For any additional context regarding the dataset, feel free to revisit HW02.

```python
bus = pd.read_csv('data/bus.csv', encoding='ISO-8859-1').rename(columns={"business id column": "bid"})
bus['postal5'] = bus['postal_code'].str[:5]
ins = pd.read_csv('data/ins.csv')
ins['timestamp'] = pd.to_datetime(ins['date'], format='%m/%d/%Y %I:%M:%S %p')
ins['bid'] = ins['iid'].str.split("_", expand=True)[0].astype(int) 

ins = ins[ins["score"] > 0]
```

```python
bus.head()
```

We also join the `ins` `DataFrame` with the appropriate portion of the `bus` `DataFrame` as was done in HW03.

```python
ins_named = pd.merge(ins, bus[['bid', 'name', 'address']])
ins_named.head()
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

# Part 1: Let Them Eat Cake

Now that you've analyzed and found out which restaurants to avoid in SF (HW03), we can turn toward the more interesting question of what dessert places are the best! For the purposes of this question, we assume that cake is the best dessert. 

<br/>

--- 

## Question 1.1

Instead of searching for cake shops, find all businesses in `ins_named` whose business `name` contains the word **`'bakery'`**, regardless of capitalization. Assign the resulting `DataFrame` to `bakery_shops`.

Your result should contain all of the original columns in `ins_named` and only the rows corresponding to bakery businesses.

**Hint:** The `.str.contains(...)` method has arguments that can make a search case-insensitive.


```python
bakery_shops = ...
bakery_shops.head()
```

```python
grader.check("q11")
```

<!-- BEGIN QUESTION -->

<br/>

--- 

## Question 1.2

Some bakery businesses appear multiple times because they were inspected more than once. Assign `frequent_bakeries` to a `DataFrame` containing **only bakery businesses that have at least 4 inspections** in `bakery_shops`.

Remember that `bid`, rather than `name`, uniquely identifies a business.

Do not aggregate the rows: `frequent_bakeries` should still contain the individual inspection records for qualifying businesses.

```python
bakery_counts = ...
frequent_bakeries = ...
frequent_bakeries.head()
```

```python
grader.check("q12")
```

<!-- END QUESTION -->

<br/>

--- 
## Question 1.3

Create a `DataFrame` named `latest_bakery_inspections` containing **exactly one row for each `bid` in `frequent_bakeries`**. For each business, keep the inspection with the **most recent `timestamp`**.

Keep the original columns and original row indices from `frequent_bakeries`. The order of rows in your final `DataFrame` does not matter.

**Hint:** One approach is to use `groupby(...)` together with `idxmax()` and then select those rows with `.loc`.

```python
latest_idx = ...
latest_bakery_inspections = ...
latest_bakery_inspections.head()
```

```python
grader.check("q13")
```

<br/>

---

## Question 1.4

Using `frequent_bakeries`, identify bakery businesses whose inspection performance has been **consistently strong**.

For each `bid`, compute:

- the **mean inspection score**, and
- the **minimum inspection score**.

Create a `DataFrame` named `consistent_bakeries` that contains only businesses satisfying **both** of the following conditions:

1. their mean inspection score is at least **90**, and  
2. their minimum inspection score is at least **85**.

Your final `DataFrame` should:

- use `bid` as the index,
- contain exactly two columns named `mean_score` and `min_score`, in that order, and
- be sorted by `mean_score` from highest to lowest.

Do **not** reset the index.

**Hint:** You may find `groupby(...)` together with `.agg(...)` useful.

```python
consistent_bakeries = ...

consistent_bakeries.head()
```

```python
grader.check("q14")
```

<br/>

---
## Question 1.5

Using `frequent_bakeries`, create a `Series` named `bakery_score_ranges` containing the **inspection-score range** for every bakery business, where the range is defined as:

`maximum score - minimum score`

The index should be `bid`, and the values should be sorted from the **largest range to the smallest range**.

**Hint:** You may find `.agg(['max', 'min'])` useful, although other pandas approaches are also valid.

```python
bakery_min_max = ...
bakery_score_ranges = ...
bakery_score_ranges.head()
```

```python
grader.check("q15")
```

<br/><br/>

--- 

## Question 2.0

As a final challenge, use **all restaurant inspections in `ins_named`**, not just bakeries.

A restaurant location is identified by the combination of its `name` and `address`. Among restaurant locations with **at least 4 inspections**, find the location whose inspection scores are the **most consistent**, where consistency is measured by the **smallest standard deviation** of `score`.

Assign `most_consistent_location` to the **business name as a string**.

To make ties deterministic, if multiple locations have the same standard deviation, choose the one that comes first alphabetically by `name`; if those names are also tied, choose the one that comes first alphabetically by `address`.

You may create temporary `DataFrame`s or `Series` objects as needed. Use pandas operations rather than loops.

```python
location_stats = ...
eligible_locations = ...
most_consistent_location = ...
most_consistent_location
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

## Question 2.1 — Interpretation

Examine the `consistent_bakeries` DataFrame you created in Question 1.4.

Without writing additional code, answer the following in **3–5 sentences**:

1. Which bakery (`bid`) shown in the first five rows has the **strongest overall inspection performance**? Justify your answer using both its `mean_score` and `min_score`.
2. Compare bakeries `3186` and `67049`. What do their mean and minimum scores tell you about their relative inspection performance?
3. Briefly explain why examining both `mean_score` and `min_score` is more informative than considering only `mean_score`.

Type your answer here, replacing this text

<!-- END QUESTION -->

## Congratulations! You have finished Homework 04! ##

### Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers.  If you run into any issues when running this cell, feel free to check the [Debugging Guide](https://mtu.instructure.com/courses/1619791/pages/debugging-guide).


### Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q11 (points: [2, 2, 2])
```
>>> type(bakery_shops) == pd.DataFrame
True
```
```
>>> list(bakery_shops.columns) == list(ins_named.columns)
True
```
```
>>> bakery_shops['name'].str.contains('bakery', case=False, na=False).all()
np.True_
```


Tests for q12 (points: [3, 3, 3])
```
>>> type(frequent_bakeries) == pd.DataFrame
True
```
```
>>> list(frequent_bakeries.columns) == list(bakery_shops.columns)
True
```
```
>>> frequent_bakeries['bid'].value_counts().min() >= 4 if len(frequent_bakeries) > 0 else True
np.True_
```


Tests for q13 (points: [3, 3, 4])
```
>>> type(latest_bakery_inspections) == pd.DataFrame
True
```
```
>>> latest_bakery_inspections['bid'].is_unique
True
```
```
>>> set(latest_bakery_inspections['bid']) == set(frequent_bakeries['bid'])
True
```


Tests for q14 (points: [3, 3, 4])
```
>>> isinstance(consistent_bakeries, pd.DataFrame)
True
```
```
>>> list(consistent_bakeries.columns) == ['mean_score', 'min_score']
True
```
```
>>> consistent_bakeries.index.name == 'bid' and (consistent_bakeries['mean_score'] >= 90).all() and (consistent_bakeries['min_score'] >= 85).all() and consistent_bakeries['mean_score'].is_monotonic_decreasing
True
```


Tests for q15 (points: [2, 3, 3])
```
>>> type(bakery_score_ranges) == pd.Series
True
```
```
>>> bakery_score_ranges.index.name == 'bid'
True
```
```
>>> (bakery_score_ranges >= 0).all()
np.True_
```


Tests for q2 (points: [4, 4, 4])
```
>>> isinstance(most_consistent_location, str)
True
```
```
>>> most_consistent_location in set(ins_named['name'])
True
```
```
>>> _q2_counts = ins_named.groupby(['name', 'address']).size()
>>> _q2_eligible = _q2_counts[_q2_counts >= 4]
>>> len(_q2_eligible) > 0
True
```
