```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

# Homework 03: Food Safety (Continued)

You must submit this assignment to Gradescope by the on-time deadline. **We strongly encourage you to plan to submit your work to Gradescope several days (hours) before the stated deadline.** This way, you will have ample time to reach out to staff for support if you encounter difficulties with submission. While course staff is happy to help guide you with submitting your assignment ahead of the deadline, we will not respond to last-minute requests for assistance (TAs need to sleep, after all!).

Please read the instructions carefully when you are submitting your work to Gradescope.




## This Assignment

In this homework, we will continue our exploration of restaurant food safety scores for restaurants in San Francisco. The main goal for this assignment is to focus more on the analysis of the dataset, building on the data cleaning we have done earlier in HW 02. 


After this homework, you should be comfortable with:
* Reading `pandas` documentation and using `pandas` methods
* Using `.groupby` with different aggregation functions
* Chaining different `pandas` functions and methods to find answers to exploratory questions



## Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.

Finally, unless we state otherwise, **do not use for loops or list comprehensions**. The majority of this assignment can be done using built-in commands in `pandas` and `NumPy`.  Our autograder isn't smart enough to check, but you're depriving yourself of key learning objectives if you write loops / comprehensions, and you also won't be ready for the midterm.

### Debugging Guide
If you run into any technical issues, we highly recommend checking out the [Debugging Guide](https://mtu.instructure.com/courses/1619791/pages/debugging-guide). In this guide, you can find general questions about Jupyter notebooks / Datahub, Gradescope, and common pandas errors.

```python
import numpy as np
import pandas as pd

import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
sns.set()
plt.style.use('fivethirtyeight')
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

In HW02, we took you through the entire process of reading data from a file to perform some exploration of the data. Here, we again load the dataset that we will be using in HW03 along with some of the columns we had added in HW02. For any additional context regarding the dataset, feel free to revisit HW02.

```python
bus = pd.read_csv('data/bus.csv', encoding='ISO-8859-1').rename(columns={"business id column": "bid"})
bus['postal5'] = bus['postal_code'].str[:5]
ins = pd.read_csv('data/ins.csv')
ins['timestamp'] = pd.to_datetime(ins['date'], format='%m/%d/%Y %I:%M:%S %p')
ins['bid'] = ins['iid'].str.split("_", expand=True)[0].astype(int) 

# This code is essential for the autograder to function properly. Do not edit.
ins_test = ins
```

```python
bus.head()
```

# Part 1: Zip-code Analysis 

Let's try and figure out whether there are any differences between names of restaurants located in even and odd ZIP codes (specifically using the 5-digit postal codes). We will break down this analysis into steps with the end goal of figuring out the restaurant with the shortest name with an even ZIP code.

<br>

---

## Question 1.1

For each restaurant in `bus`, count how many whitespace-separated words appear in its `name`. Store the result in a pandas **Series** named `name_word_count`.

For example, a restaurant name containing three words should have a value of `3`. **Do not use for loops or list comprehensions.**

```python
name_word_count = ...
name_word_count

```

```python
grader.check("q11")
```

<br>

---

## Question 1.2

The `postal5` column may contain ZIP codes that are not in the official San Francisco ZIP-code list. The code below loads the valid ZIP codes.

Create a new `DataFrame` named `bus_valid` that is a is a [`copy` of `bus`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.copy.html#pandas.DataFrame.copy) containing only rows whose `postal5` value appears in `valid_zips`. For the rest of Part 1, use `bus_valid`.

```python
valid_zips = pd.read_json('data/sf_zipcodes.json')['zip_codes']
valid_zips = valid_zips.astype("string")
```

```python
bus_valid = ...
bus_valid.head()

```

```python
grader.check("q12")
```

<br>

---

## Question 1.3

Using `bus_valid`, compute the **average restaurant-name length (in characters)** for each 5-digit ZIP code. Store the result in a pandas **Series** named `zip_avg_name_length`.

The index should contain the ZIP codes, the values should contain the corresponding averages, and the index should be sorted in ascending order. **Do not use loops or list comprehensions.**

```python
zip_avg_name_length = ...
zip_avg_name_length

```

```python
grader.check("q13")
```

<br>

---

## Question 1.4

Using `zip_avg_name_length`, find the ZIP code whose restaurants have the **shortest average name length**. Store that ZIP code as a **string** in `shortest_avg_name_zip`.

If multiple ZIP codes have the same minimum average, break the tie by choosing the numerically smallest ZIP code.

Note: [sort_values documentation](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sort_values.html) to see how you can sort by multiple values.

```python
shortest_avg_name_zip = ...
shortest_avg_name_zip

```

```python
grader.check("q14")
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

# Part 2: Inspecting the Inspections 

<br>

---

## Question 2.1

The `timestamp` column in `ins` contains each inspection date as a pandas datetime value. Create a pandas **Series** named `inspection_year` containing only the four-digit year for every inspection.

Do not modify `ins` in this question.

```python
inspection_year = ...
inspection_year

```

```python
grader.check("q21")
```

<br>

---

## Question 2.2

Count how many inspections occurred for every combination of **year** and inspection `type`.

Create a `DataFrame` named `ins_year_type_counts` with a two-level index `['year', 'type']` and one column named `'Count'`. Use `.groupby`; do not use loops or list comprehensions.

```python
ins_year_type_counts = ...
ins_year_type_counts

```

```python
grader.check("q22")
```

<!-- BEGIN QUESTION -->

<br>

---

## Question 2.3

Reshape the inspection counts into a more readable table. Create a `DataFrame` named `inspection_type_by_year` where:

- each row is an inspection `type`,
- each year is a separate column,
- each value is the number of inspections of that type in that year,
- missing year/type combinations are filled with `0`, and
- a final column named `'Total'` contains the total number of inspections for that type.

Sort the rows by `'Total'` from largest to smallest. If totals tie, preserve the alphabetical order of the inspection types produced by the pivot table.

```python
inspection_type_by_year = ...

...
inspection_type_by_year

```

```python
grader.check("q23")
```

<!-- END QUESTION -->

Notice that inspection scores appear only to be assigned to `Routine - Unscheduled` inspections. It is reasonable for inspection types such as `New Ownership` and `Complaint` to have no associated inspection scores, but we might be curious why there are no inspection scores for the `Reinspection/Followup` inspection type.

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

# Part 3: Joining Data Across Tables

In this question, we will start to connect data across multiple tables. We will be using the `merge` function. 

<br>

---

## Question 3.1

For the remainder of Part 3, we only want inspections with an actual positive score.

First create `scored_ins`, a copy of the rows in `ins` where `score > 0`. Then merge it with the appropriate columns of `bus` to create `ins_named`. The resulting `ins_named` should contain every column from `scored_ins` plus the restaurant `name` and `postal5`.

Use `bid` as the merge key and keep every row from `scored_ins`.

```python
scored_ins = ...
ins_named = ...
ins_named.head()

```

```python
grader.check("q31")
```

<!-- BEGIN QUESTION -->

<br>

---

## Question 3.2

Using `ins_named`, summarize inspection performance for each business. Create a `DataFrame` named `ten_lowest_average` containing the **10 businesses with the lowest average inspection score**.

For each `bid`, include:

- `name`: the restaurant name,
- `inspection_count`: number of scored inspections, and
- `average_score`: mean inspection score.

Use `bid` as the index. Sort by `average_score` ascending, then by `name` ascending to break ties, and keep only the first 10 rows. **Do not use loops or list comprehensions.**

**Helpful documentation:**
- [`DataFrameGroupBy.aggregate`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.api.typing.DataFrameGroupBy.aggregate.html)
- [`DataFrame.sort_values`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sort_values.html)
- [`DataFrame.rename`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.rename.html)

```python
ten_lowest_average = ...

# DO NOT USE LIST COMPREHENSIONS OR LOOPS OF ANY KIND!!!

...
ten_lowest_average

```

```python
grader.check("q32")
```

<!-- END QUESTION -->

<br>

---

## Question 3.3

Which business received the **largest number of scored inspections**?

Using the business-level summary from Question 3.2, assign `most_inspected_bid` to the `bid` of that business. If multiple businesses have the same inspection count, break the tie by restaurant `name` in alphabetical order.

Store the final answer as a Python `int`.

```python
most_inspected_bid = ...
most_inspected_bid

```

```python
grader.check("q33")
```

<!-- BEGIN QUESTION -->

<br>

---

## Question 3.4

For the business identified in Question 3.3, compare its **earliest** and **latest** scored inspections.

Assign `latest_score_higher` to `True` if its latest inspection score is strictly higher than its earliest inspection score; otherwise assign `False`.

Use the `timestamp` column to determine the inspection order.

```python
latest_score_higher = ...
latest_score_higher
```

```python
grader.check("q34")
```

<!-- END QUESTION -->

## Summary of Inspections Data

We have done a lot in this homework! 
 
- Broke down the inspection scores in detail using `.groupby` and `pivot_table`
- Joined the business and inspection data and identified the name of the restaurant with the worst rating


Over the course of the last 2 homeworks, we hope you have become more familiar with `pandas` - in terms of identifying when to use particular functions, how they work, when they can support EDA - as well as with EDA and Data Cleaning, as part of the broader Data Science Lifecycle. These tools will serve you well as a data scientist!

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Congratulations! You have finished Homework 03! ##

### Submission Instructions

Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers. Once you submit this file to the HW03 Coding assignment on Gradescope, Gradescope will automatically submit a PDF file with your written answers to the HW03 Written assignment. If you run into any issues when running this cell, feel free to check the [Debugging Guide](https://mtu.instructure.com/courses/1619791/pages/debugging-guide).

**Important**: Please check that your written responses were generated and submitted correctly to the HW03 Written Assignment.

**You are responsible for ensuring your submission follows our requirements and that the PDF for HW03 written answers was generated/submitted correctly. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline. 

### Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q11 (points: [1, 1])
```
>>> type(name_word_count) == pd.Series
True
```
```
>>> name_word_count.shape == (len(bus),)
True
```


Tests for q12 (points: [3])
```
>>> type(bus_valid) == pd.DataFrame and bus_valid['postal5'].isin(valid_zips).all()
np.True_
```


Tests for q13 (points: [1, 1])
```
>>> type(zip_avg_name_length) == pd.Series
True
```
```
>>> zip_avg_name_length.index.is_monotonic_increasing
True
```


Tests for q14 (points: [1, 3])
```
>>> type(shortest_avg_name_zip) == str
True
```
```
>>> shortest_avg_name_zip in set(bus_valid['postal5'].astype(str))
True
```


Tests for q21 (points: [1, 1, 2])
```
>>> type(inspection_year) == pd.Series
True
```
```
>>> len(inspection_year) == len(ins)
True
```
```
>>> inspection_year.index.equals(ins.index)
True
```


Tests for q22 (points: [1, 1, 2])
```
>>> type(ins_year_type_counts) == pd.DataFrame
True
```
```
>>> list(ins_year_type_counts.columns) == ['Count']
True
```
```
>>> isinstance(ins_year_type_counts.index, pd.MultiIndex)
True
```


Tests for q23 (points: [2, 2])
```
>>> type(inspection_type_by_year) == pd.DataFrame
True
```
```
>>> 'Total' in inspection_type_by_year.columns
True
```


Tests for q31 (points: [1, 1])
```
>>> all(ins_named['score'] > 0)
True
```
```
>>> 'name' in ins_named.columns and 'postal5' in ins_named.columns
True
```


Tests for q32 (points: [1, 1, 2, 2])
```
>>> type(ten_lowest_average) == pd.DataFrame
True
```
```
>>> ten_lowest_average.shape == (10, 3)
True
```
```
>>> list(ten_lowest_average.columns) == ['name', 'inspection_count', 'average_score']
True
```
```
>>> ten_lowest_average.index.name == 'bid'
True
```


Tests for q33 (points: [1, 1])
```
>>> type(most_inspected_bid) == int
True
```
```
>>> most_inspected_bid in set(ins_named['bid'])
True
```


Tests for q34 (points: [1])
```
>>> latest_score_higher in [True, False]
True
```
