```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Lab 07: Visualization, Transformations, and KDEs
**DATA 2201 — Fall 2026 | 100 points**

In this lab, you will explore World Bank income and adult literacy data for countries and territories around the world. You will create and compare visualizations, examine how transformations change the appearance of a relationship, and explore how kernel density estimation (KDE) provides a smooth view of a distribution.

### Learning Objectives
By the end of this lab, you should be able to:
- Choose appropriate plots for quantitative data.
- Create histograms, rug plots, density curves, and scatter plots with clear titles and axis labels.
- Distinguish between count and density when interpreting distributions.
- Explore how transformations and bandwidth choices affect visual patterns.
- Construct normalized Gaussian kernels and implement a simple KDE function.

### Tasks and Grading
- **Part 1 — Visualization: 60 points.** Explain plot choices and visualize income and literacy distributions and their relationship.
- **Part 2 — Kernel Density Estimation: 40 points.** Plot normalized kernel contributions and implement a KDE that handles missing observations.

**Dataset:** `data/world_bank_misc.csv`

### Before Submitting
Restart the kernel and run all cells in order. Confirm that the available tests pass, check your written responses, and save the notebook with all plot outputs visible before submitting to Gradescope.

## Setup

Note that we configure a custom default figure size. Virtually every default aspect of matplotlib [can be customized](https://matplotlib.org/users/customizing.html).

```python
import pandas as pd
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt

plt.style.use('fivethirtyeight') # Use plt.style.available to see more styles
sns.set()
sns.set_context("talk")
plt.rcParams['figure.figsize'] = (8, 5)
%matplotlib inline
```

### Get the Data

Let us load some World Bank data into a `pd.DataFrame` object named ```wb```.

```python
wb = pd.read_csv("data/world_bank_misc.csv", index_col=0)
wb.head()
```

This table contains some interesting columns.  Take a look:

```python
list(wb.columns)
```

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />


## Part 1: Scaling 



In the first part of this assignment we will look at the distribution of values for combined adult literacy rate as well as the gross national income per capita. The code below creates a copy of the DataFrame that contains only the two Series we want, and then drops all rows that contain null values in either column.

**Note:** For this lab we are dropping null values without investigating them further. However, this is generally not the best practice and can severely affect our analyses.

Here the combined literacy rate is the sum of the female and male literacy rates as reported by the World Bank. 0 represents no literacy, and 200 would represent total literacy by both genders that are included in the World Bank's dataset.

In this lab, we will be using the `sns.histplot`, `sns.rugplot`, and `sns.displot` function to visualize distributions. You may find it useful to consult the seaborn documentation on [distributions](https://seaborn.pydata.org/tutorial/distributions.html) and [functions](https://seaborn.pydata.org/tutorial/function_overview.html) for more details.

```python
#creates a DataFrame with the appropriate index
df = pd.DataFrame(index=wb.index)

#copies the Series we want
df['lit'] = wb['Adult literacy rate: Female: % ages 15 and older: 2005-14'] + \
  wb["Adult literacy rate: Male: % ages 15 and older: 2005-14"]
df['inc'] = wb['Gross national income per capita, Atlas method: $: 2016']

#the line below drops all records that have a NaN value in either column
df.dropna(inplace=True)
print("Original records:", len(wb))
print("Final records:", len(df))
```

```python
df.head(5)
```

<br>

--- 

### Question 1

Suppose we wanted to build a histogram of our data to understand the distribution of literacy rates and income per capita individually. We can use [seaborn's `countplot`](https://seaborn.pydata.org/generated/seaborn.countplot.html) to create bar charts from categorical data. 

```python
sns.countplot(x = "lit", data = df)
plt.xlabel("Combined literacy rate: % ages 15 and older: 2005-14")
plt.title('World Bank Combined Adult Literacy Rate')
```

```python
sns.countplot(x = "inc", data = df)
plt.xlabel('Gross national income per capita, Atlas method: $: 2016')
plt.title('World Bank Gross National Income Per Capita')
```

<!-- BEGIN QUESTION -->

### Question 1 — Choosing a distribution plot (10 points)
The two count plots above treat each distinct numeric value as a separate category. In **2–3 sentences**, explain how a histogram would organize these values differently and why that makes it more useful for examining income and literacy distributions.

**Hint:** Think about individual values versus ranges of values.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

---

### Question 2 — Density histograms (15 points)
Create two vertically stacked **density histograms** using the prepared `df`:
- Top: combined literacy (`lit`).
- Bottom: income per capita (`inc`).

Use **20 bins** in each plot. Add titles, label both axes with appropriate units or scales, and use a readable figure layout. The y-axes should show **density**, rather than count.

**Hint:** You can use `plt.subplots` and `sns.histplot`. Consider the `stat` and `ax` parameters.

```python
# Create the two density histograms here.
fig, axes = ...
sns.  
sns.
axes[0].
axes[1]
...
...
```

<!-- END QUESTION -->

<br>

In the cell below, we explore overlaying a rug plot on top of a histogram using `rugplot`. Note that the rug plot is hard to see.

```python
sns.histplot(x="inc", data = df)
sns.rugplot(x="inc", data = df)
plt.xlabel('Gross national income per capita, Atlas method: $: 2016')
plt.title('World Bank Gross National Income Per Capita')
```

One way to make it easier to see the difference between the rug plot and the bars is to set a different color, for example:

```python
sns.histplot(x="inc", data = df, color = "lightsteelblue")
sns.rugplot(x="inc", data = df)
plt.xlabel('Gross national income per capita, Atlas method: $: 2016')
plt.title('World Bank Gross National Income Per Capita')
```

There is also another function called `kdeplot` which plots a Kernel Density Estimate as described in class, and covered in more detail later in this lab.



<!-- BEGIN QUESTION -->

<br>

---

### Question 3 — Literacy distribution with rug and KDE (13 points)
Use `sns.displot` to visualize **combined literacy (`lit`)** with a density histogram, an overlaid KDE curve, and rug marks. Use **12 histogram bins** and make the figure approximately **8 inches wide and 5 inches tall**. Add a descriptive title and label the axes.

**Hint:** Explore the `kde`, `rug`, `stat`, `height`, and `aspect` parameters.

```python
# Create the literacy distribution plot here.
g = sns.  
g.set
g.ax.
...
```

The bars show a density histogram, the smooth curve is a KDE, and the rug marks show individual observations. These layers provide complementary views of the same literacy data.

<!-- END QUESTION -->

In the figure above, the y-axis is labeled by the counts. We can also label the y-axis by the density. An example is given below, this time using the literacy data from the beginning of this lab.

```python
sns.displot(x="lit", data = df, rug = True, kde = True, stat = "density")
plt.xlabel("Adult literacy rate: Combined: % ages 15 and older: 2005-14", fontsize=12)
plt.title('World Bank Combined Adult Literacy Rate')
```

Observations:
* You'll also see that the y-axis value is no longer the count. Instead it is a value such that the total **area** in the histogram is 1. For example, the area of the last bar is approximately 22.22 * 0.028 = 0.62

* The KDE is a smooth estimate of the distribution of the given variable. The area under the KDE is also 1. While it is not obvious from the figure, some of the area under the KDE is beyond the 100% literacy. In other words, the KDE is non-zero for values greater than 100%. This, of course, makes no physical sense. Nonetheless, it is a mathematical feature of the KDE.

We'll talk more about KDEs later in this lab.

<!-- BEGIN QUESTION -->

<br>

---

### Question 4 — Income and literacy relationship (12 points)
Create a scatter plot using **income (`inc`) on the x-axis** and **combined literacy (`lit`) on the y-axis**. Use small, partially transparent markers so overlapping observations are easier to see. Add a title and axis labels with the correct units or scale. Use the original values without transformations or a fitted line.

**Hint:** You can use `sns.scatterplot`; consider `s` and `alpha`.

```python
# Create the scatter plot here.
plt.figure(...)
sns.
...
...
...
...
...
```

<!-- END QUESTION -->

<br> 

We can better assess the relationship between two variables when they have been straightened because it is easier for us to recognize linearity.

In the cell below, we see a scatter plot of log-transformed income per capita against literacy rate.


```python
sns.scatterplot(x = df['lit'], y = np.log10(df['inc']))
plt.xlabel("Adult literacy rate: Combined: % ages 15 and older")
plt.ylabel('Gross national income per capita (log scale)')
plt.title('World Bank: Gross National Income Per Capita vs\n Combined Adult Literacy Rate');
```

### Observations:
- This scatter plot looks better. The relationship is closer to linear.
- We can think of the log-linear relationship between x and y, as follows: a constant change in x corresponds to a percent (scaled) change in y.
- Choosing the best transformation for a relationship is often a balance between keeping the model simple and straightening the scatter plot.

### We’ll discuss data transformation using log10 in Lecture 14: Visualization III.

<br/><br/>
<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Part 2: Kernel Density Estimation 

In this part of the lab you will develop a deeper understanding of how kernel density estimation works.
- Explain KDE briefly within the lab

### Overview

Kernel density estimation is used to estimate a probability density function (i.e. a density curve) from a set of data. Just like a histogram, a density function's total area must sum to 1.

KDE centrally revolves around this idea of a "kernel". A kernel is a function whose area sums to 1. The three steps involved in building a kernel density estimate are:
1. Placing a kernel at each observation
2. Normalizing kernels so that the sum of their areas is 1
3. Summing all kernels together

The end result is a function, that takes in some value `x` and returns a density estimate at the point `x`.

When constructing a KDE, there are several choices to make regarding the kernel. Specifically, we need to choose the function we want to use as our kernel, as well as a bandwidth parameter, which tells us how wide or narrow each kernel should be. We will explore these ideas now.

```python
data3pts = np.array([2, 4, 9])
sns.displot(data3pts, kde = True, stat = "density");
```

To understand how KDEs are computed, we need to see the KDE outside the given range. 

The easiest way to do this was to use an old function called `distplot`. But, `distplot` is now deprecated, and it will be removed at a future date. If you get an error that says that `distplot` is not a valid function, sorry, you are too far in the future to do this code cell.

Using a deprecated function will often result in a `UserWarning`. You may follow their suggestions to adapt your code. In this case, we can use `displot` [(documentation)](https://seaborn.pydata.org/archive/0.11/generated/seaborn.displot.html#seaborn.displot) with some additional parameters. These additional parameters are needed because the default values of `displot` are different from `distplot`; we can manually set them to be the same.

```python
# Run this cell to create KDE plot and histogram using displot; no further action is needed.
plt.figure(figsize=(5, 15))
sns.displot(data3pts, kde=True, stat="density", kde_kws={"cut":4}, bins=2, 
            height=5, aspect=1.5);
```

One question you might be wondering is how the kernel density estimator decides how "wide" each point should be. It turns out this is a parameter you can set called `bw`, which stands for bandwidth. For example, the code below gives a bandwidth value of 0.5 to each data point. You'll see the resulting KDE is quite different. Try experimenting with different values of bandwidth and see what happens.

```python
# Run this cell to plot displot with the specified bandwidth. 
plt.figure(figsize=(5, 15))
sns.displot(data3pts, kde=True, stat="density", 
            kde_kws={"bw_method": 0.5, "cut":4}, bins=2, height=5, aspect=1.5);
```

<br>

---

As mentioned above, the kernel density estimate is just the sum of a bunch of copies of the kernel, each centered on our data points. The default kernel used by the `displot` function (as well as `kdeplot`) is the Gaussian kernel, given by:

$$\Large
K_\alpha(x, z) = \frac{1}{\sqrt{2 \pi \alpha^2}} \exp\left(-\frac{(x - z)^2}{2  \alpha ^2} \right)
$$

We've implemented the Gaussian kernel for you in Python below. Here, `alpha` is the smoothing or bandwidth parameter $\alpha$ for the KDE, `z` is the center of the Gaussian (i.e., a data point or an array of data points), and `x` is an array of values of the variable whose distribution we are plotting. In other words, `z` represents the center point of our smooth KDE bell curve, while `x` represents the range of values over which we want to generate the KDE plot.

```python
def gaussian_kernel(alpha, x, z):
    """
    Compute the Gaussian density estimate for values in x.

    Args:
        alpha: the smoothing parameter to pass to the kernel.
        x: an array of values whose density will be calculated.
        z: center of Gaussian.

    Returns:
        The smoothed estimate at values of x.
    """    
    return 1.0/np.sqrt(2. * np.pi * alpha**2) * np.exp(-(x - z) ** 2 / (2.0 * alpha**2))
```

For example, we can plot the Gaussian kernel centered at 9 with $\alpha$ = 0.5 as below: 

```python
xs = np.linspace(-2, 12, 200)
alpha = 0.5
kde_curve = [gaussian_kernel(alpha, x, 9) for x in xs]
plt.plot(xs, kde_curve);
```

<!-- BEGIN QUESTION -->

### Question 5 — Normalized kernels and their sum (20 points)
Using the supplied `gaussian_kernel`, plot the Gaussian contributions centered at **2, 4, and 9** on the same axis with **alpha = 1**. Divide each contribution by **3** so its area is one-third.

Add the **sum of the three contributions** as a thicker line. Include a legend identifying the three contributions and the total KDE, a title, and axes labeled **Value** and **Density**.

**Hint:** Use the existing grid `xs` or create one with `np.linspace`. You can compute the three contributions separately and add their arrays; no loop is necessary.

```python
# Plot the contributions and their sum here.
...
```

<!-- END QUESTION -->

The next supplied example uses alpha = 0.5 instead of 1. Run it to see how the same three observations produce a different density curve when the bandwidth changes.

```python
# Run this cell to plot the sum of the kernels; no further action is needed.
xs = np.linspace(-2, 12, 200)
alpha = 0.5
kde_curve = np.array([1/3 * gaussian_kernel(alpha, x, data3pts) for x in xs])
plt.plot(xs, np.sum(kde_curve, axis=1));
```

Recall that earlier we plotted the kernel density estimation for the logarithm of the income data, as shown again below.

```python
# Run this cell to plot KDE of log income; no further action is needed.
ax = sns.displot(data=df, x=np.log10(df['inc']), kind='kde', rug=True, aspect=1.4)
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016', fontsize=13)
plt.title('World Bank Gross National Income Per Capita');
```

In the cell below, a similar plot is shown using what was done above. Try out different values of alpha in {0.1, 0.2, 0.3, 0.4, 0.5}. You will see that when `alpha=0.2`, the graph matches the previous graph well, except that the `displot` function hides the KDE values outside the range of the available data. Recall that `alpha` represents the spread of each individual kernel curve, which in turn affects the smoothness of the final KDE plot. When would you expect a smoother plot - when alpha = 0.1 or 0.5?

```python
# Run this cell to plot KDE of log income; try out different bandwidths!
xs = np.linspace(1, 6, 200)
alpha = 0.5
kde_curve = np.array([1/len(df['inc']) * gaussian_kernel(alpha, x, np.log10(df['inc'])) for x in xs])
plt.xlabel('Log Gross national income per capita, Atlas method: $: 2016', fontsize=13)
plt.ylabel("Density")
plt.title('World Bank Gross National Income Per Capita')
plt.plot(xs, np.sum(kde_curve, axis = 1));
```

<!-- BEGIN QUESTION -->

### Question 6: KDE Bandwidth — Interpret the Plot (10 points)

Run the supplied code to compare two KDE curves for combined literacy.

In 2–3 sentences:
1. Identify which curve is smoother.
2. Explain how increasing the bandwidth changes the detail visible in the distribution.

```python
plt.figure(figsize=(9, 5))
sns.kdeplot(data=df, x='lit', bw_adjust=0.5,
            color='royalblue', label='bw_adjust = 0.5')
sns.kdeplot(data=df, x='lit', bw_adjust=2,
            color='darkorange', label='bw_adjust = 2')
plt.xlabel('Combined literacy (0–200)')
plt.ylabel('Density')
plt.title('Combined literacy: effect of KDE bandwidth')
plt.legend()
plt.tight_layout()
plt.show()
```

Type your answer here...

<!-- END QUESTION -->

<br>

---

### Question 7 — A KDE that handles missing observations (20 points)
Implement `kde(kernel, alpha, x, data)` to return the KDE at a **single query value** `x`. Before computing the density, exclude NaN and positive/negative infinity from `data`. Normalize by the number of **remaining finite observations**.

$$f_\alpha(x)=\frac{1}{n}\sum_{i=1}^{n}K_\alpha(x,z_i).$$

Use the kernel function passed through `kernel`, rather than hard-coding `gaussian_kernel`. Assume alpha is positive and at least one finite observation remains. Do not modify the input.

**Hint:** Convert `data` to a NumPy array, use `np.isfinite` to select valid values, and average the kernel values. No loop is required.

```python
def kde(kernel, alpha, x, data):
    values = ...
    finite = ...
    density= ...
    return density
```

```python
grader.check("q7")
```

<br><br>

<hr style="border: 5px solid #8a8c8c;" />
<hr style="border: 1px solid #ffcd00;" />

## Congratulations! You have finished Lab 07!


If you run into any issues when running this cell, feel free to check the [Debugging Guide](https://mtu.instructure.com/courses/1619791/pages/debugging-guide).

**Important**: Please check that your written responses were generated and submitted correctly.

If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q7 (points: [4, 4])
```
>>> np.isclose(kde(gaussian_kernel, 1, 2, [1, 2, 3]), np.mean(gaussian_kernel(1, 2, np.array([1.0, 2.0, 3.0]))))
np.True_
```
```
>>> np.isclose(kde(gaussian_kernel, 0.5, 4, [2, np.nan, 4, 9, np.inf, -np.inf]), np.mean(gaussian_kernel(0.5, 4, np.array([2.0, 4.0, 9.0]))))
np.True_
```
